### CS 210: Data Management for Data Science
#### Fall 2023

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Arnav Dashaputra"
COLLABORATORS = ""
```

---

# **CS 210**: HW2



Read the problem descriptions carefully. If you have questions about any of the problems in this assignment, please post them so other students can see.  

### Question 1
Write a program that reads words from a **filename**, which is given as
a string argument. It should **return** the words from the file in a list, sorted in reverse alphabetical order (case insensitive) .

For instance, if the file has

```
bell tea  Zebra
apple
yellow
```

Then the output should be

`['Zebra', 'yellow', 'tea', 'bell', 'apple']`

<!-- grade_id: cell-91029b5f4c39aa00, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE

def reverse_sorted_words(filename):
    try:
        with open(filename, 'r') as file:
            
            content = file.read()
            words = content.split()
            words = [word.lower() for word in words]
            words.sort(reverse=True)
            
            return words
        
    except FileNotFoundError:
        print(f"'{filename}' not found.")
        return []

# raise NotImplementedError()
```

Example output for 

`reverse_sorted_words("words.txt")`

```
['Zebra', 'whispering', 'visitor', 'utopian', 'thrill', 'ray', 'point', 'Pink', 'pan', 'launch', 'hurry', 'houses', 'frog', 'fang', 'continue', 'buzz', 'Beach', 'ask', 'Apple', 'airport', 'advise', 'activity']
```

<!-- grade_id: cell-d994a1dde895f631, points: 0, grade: true, solution: false -->
```python
# Use this code block to test your program, please make your own tests too.

answer_key = ['whispering', 'visitor', 'utopian', 'thrill', 'tangible', 'ray', 'point', 'pan', 'launch', 'hurry', 'houses', 'frog', 'five', 'fang', 'continue', 'buzz', 'ask', 'airport', 'advise', 'activity']
rev_words = reverse_sorted_words("words.txt")
assert len(answer_key) == len(rev_words)
for ans, res in zip(answer_key, rev_words):
    assert ans == res
    
answer_key2 = ['whispering', 'visitor', 'visitor', 'visitor', 'utopian', 'thrill', 'tangible', 'ray', 'ray', 'point', 'pan', 'launch', 'hurry', 'hurry', 'hurry', 'houses', 'houses', 'houses', 'frog', 'frog', 'five', 'five', 'fang', 'continue', 'continue', 'buzz', 'buzz', 'buzz', 'buzz', 'buzz', 'ask', 'airport', 'advise', 'advise', 'activity']
rev_words2 = reverse_sorted_words("words2.txt")
assert len(answer_key2) == len(rev_words2)
for ans2, res2 in zip(answer_key2, rev_words2):
    assert ans2 == res2
```

### Question 2

Your professor asked your TA to collect student information and store it in a file called **studentInfo.txt**, but this TA spilled Pepsi on the keyboard while entering the data. As a result, there are extra or missing spaces and missing commas throughout the file. For example, the file has

```
wd50,WillDavidson, 090,3
hpj4332, Helen              Jordan1983.88
xd222XavierDavies1983.65
ad12  AdDavies1983.65
```
Suppose the file contains netID, name, major, and GPA in that order. Can you help this TA make the file more readable? 

To accomplish this task, you may assume the following about each field: 
`netID`: begins with 2-3 lowercase letters followed by 1-4 digits (e.g., wd1, wxo001, tsy1234)
`name`: contains both first name and last name with their first letter capitalized (e.g., "John   Doe", "MarySmith"). There may be zero or more spaces between the first name and last name. 
`major`: 3-digit code (e.g., 198)
`GPA`: floating point number between 0.0 and 4.0 (inclusive). There may or may not be a decimal point (e.g., 3.0 or 3). 

Write a function called **reformat_student_info** to convert a file into a standard format, which has exactly one comma followed by one space to separate each field. The function should accept a **filename** string parameter. Write the cleaned data to a new file called **studentInfoReformatted.csv** and format the file as follows:

`netID, name, major, GPA`

There should be exactly one comma followed by one space to separate each field. The name field should include exactly one space between first name and last name, and the GPA should have a decimal even if the GPA from the input file does not (e.g., 3.0 instead of 3). 

The reformatted file for the sample above should look like this: 

```
wd50, Will Davidson, 090, 3.0
hpj4332, Helen Jordan, 198, 3.88
xd222, Xavier Davies, 198, 3.65
ad12, Ad Davies, 198, 3.65
```
 

<!-- grade_id: cell-febddcc2fe833c4e, points: 0, grade: true, solution: true -->
```python
import csv

def reformat_student_info(filename):
    # Open the txt file
    with open(filename, 'r') as f:
        lines = f.readlines()

    records = []

    for line in lines:
        netid = ""
        first_name = ""
        last_name = ""
        class_code = ""
        gpa = ""
        
        # Extraction for netID
        idx = 0
        while idx < len(line) and not line[idx].isupper():
            netid += line[idx]
            idx += 1
        netid = netid.strip().rstrip(',')  # Remove spaces and trailing comma

        # Extraction for the first name
        while idx < len(line) and (line[idx].isalpha() or line[idx] == ' '):
            if line[idx].isupper() and first_name: # if we encounter an uppercase letter after the start of the first name
                break
            first_name += line[idx]
            idx += 1

        # Extraction for the last name
        while idx < len(line) and line[idx].isalpha():
            last_name += line[idx]
            idx += 1
        
        # Extraction for class code
        while idx < len(line) and not line[idx].isdigit():
            idx += 1
        class_code_digits = 0
        while idx < len(line) and class_code_digits < 3:
            class_code += line[idx]
            class_code_digits += 1
            idx += 1
        
        # Extraction for GPA
        while idx < len(line) and not line[idx].isdigit():
            idx += 1
        while idx < len(line) and (line[idx].isdigit() or line[idx] == '.'):
            gpa += line[idx]
            idx += 1
        
        # If GPA ends in a number with no dot, add ".0"
        if '.' not in gpa:
            gpa += ".0"

        records.append([netid, ' ' + first_name.strip() + ' ' + last_name.strip(), ' ' + class_code, ' ' + gpa])

    # Write the extracted netIDs, names, class codes, and GPA to CSV
    with open('studentInfoReformatted.csv', 'w', newline='') as csvfile:
        writer = csv.writer(csvfile)
        for record in records:
            writer.writerow(record)
```

```python
# used for testing purposes:
# import os
# os.remove("studentInfoReformatted.csv") 
```

```python
# Your cleaned file should look like this

reformat_student_info('studentInfo.txt')
f = open('studentInfoReformatted.csv')
for i, line in enumerate(f):
    print(line.strip())

f.close()
```

ed18, Edward Duran, 077, 3.5

abc1234, Mike Peterson, 123, 3.7

dw0, Danika Wilson, 083, 3.9

wd50, Will Davidson, 090, 3.0

hpj4332, Helen Jordan, 198, 3.88

xd222, Xavier Davies, 198, 3.65

ah1, Allisson Henson, 147, 2.9

ne099, Nina Estrada, 095, 2.0

sh5555, Skylar Huff, 112, 3.3

### Question 3
Write a program that reads words from a file (**filename** given
as a string parameter) and prints the occurance of each word(case insensitive). And print the words in **alphabetical** order.

For example, if the file contains text



```
Love is free
free is love
```



then the function should print

```
free:2
is:2
love:2
```



<!-- grade_id: cell-e2053a97d07532a1, points: 0, grade: true, solution: true -->
```python
def count_word(filename):

    # YOUR CODE 
    
    word_count = {}
    
    with open(filename, 'r') as file:
        for line in file:
            words = line.split()
            for i in words:
                i = i.lower() # force case
                if i in word_count:
                    word_count[i] += 1 # increment if found
                else:
                    word_count[i] = 1 # add new

    sorted_words = sorted(word_count.items())
    
    for word, count in sorted_words:
        print(f"{word}: {count}")
    
    #raise NotImplementedError()
```

```python
# Your output should look like this
count_word('words2.txt')
```

```
activity: 1
advise: 2
airport: 1
ask: 1
buzz: 5
continue: 2
fang: 1
five: 2
frog: 2
houses: 3
hurry: 3
launch: 1
pan: 1
point: 1
ray: 2
tangible: 1
thrill: 1
utopian: 1
visitor: 3
whispering: 1
```

### Question 4

Write a program that reads words passing from a string parameter and create a dictionary mapping **last** letters to the number of words **ending** with that letter.

For example, the string is `'apple banana orange grape cherry`', the program will
out put a dictionary `fruitName` that


```
print(fruitName)
```
will give us output

```
{'e': 3, 'a': 1, 'y': 1}
```







<!-- grade_id: cell-359500c2bf1aa28a, points: 0, grade: false, solution: true -->
```python
def count_last_letter(words):

    # YOUR CODE HERE
    
    counts = {}
    word_list = words.split()
    
    for word in word_list:
        last = word[-1]
        if last in counts:
            counts[last] += 1
        else:
            counts[last] = 1

    return counts
    
    # raise NotImplementedError()
```

<!-- grade_id: cell-74040366184943e7, points: 0, grade: true, solution: false -->
```python
# use this block of code to test your code, note this is just one
# testing case, please make your own tests too.

d = count_last_letter('apple banana orange grape cherry')

words = 'apple banana orange grape cherry'
fruitName = count_last_letter(words)
print(fruitName)

assert len(d) == 3
assert d['a'] == 1
```

### Question 5

We can create a numpy array from separate lists by specifying a format for each field:



```python
roster = np.zeros(4, dtype={'names': ('name', 'age', 'major', 'gpa'),
                         'formats': ('U50', 'i4', 'U4', 'f8')})

roster['name']  = ['Alice', 'Bob', 'Carol', 'Dennis']
roster['age']   = [21, 25, 18, 29]
roster['major'] = ['CS', 'Math', 'Chem', 'Phys']
roster['gpa']   = [3.8, 3.2, 4.0, 3.5]
```
where **U50** is a string of max length 50, **i4** is a 4-byte integer,
**U4** is a string of max length 4, and **f8** is an 8-byte
floating-point number. Then **roster** has value



```python
array([('Alice', 21, 'CS', 3.8), ('Bob', 25, 'Math', 3.2),
  ('Carol', 18, 'Chem', 4. ), ('Dennis', 29, 'Phys', 3.5)],
  dtype=[('name', '<U50'), ('age', '<i4'), ('major', '<U4'), ('gpa', '<f8')])
```

Then, for example, **roster['name']** will give an array of just the
student names.

Write a program that reads the numpy array **roster** we just created that has this format:

`name,age,major,gpa`


Convert this to a numpy array as shown above. Then use that array
to compute and print each of the following:



1.   The average GPA of all students
2.   The maximum GPA of students majoring in CS
3.   The number of students with a GPA over 3.5
4.   The average GPA of students who are at least 25 years old
5.   The major that has the highest average GPA among students at most 22 years old

For example, for the four students listed, your program should print:

```
3.625
3.8
2
3.35
Chem
```



<!-- grade_id: cell-6098f8cf8de55161, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE

import numpy as np

roster = np.zeros(4, dtype={'names': ('name', 'age', 'major', 'gpa'),'formats': ('U50', 'i4', 'U4', 'f8')})
roster['name']  = ['Alice', 'Bob', 'Carol', 'Dennis']
roster['age']   = [21, 25, 18, 29]
roster['major'] = ['CS', 'Math', 'Chem', 'Phys']
roster['gpa']   = [3.8, 3.2, 4.0, 3.5]


avg_gpa = np.mean(roster['gpa'])
max_cs_gpa = np.max(roster['gpa'][roster['major'] == 'CS'])
top_gpa = np.sum(roster['gpa'] > 3.5)
older = np.mean(roster['gpa'][roster['age'] >= 25])
majors_22_below = roster['major'][roster['age'] <= 22]
gpa_22_below = roster['gpa'][roster['age'] <= 22]
uniques = np.unique(majors_22_below)

max_avg_gpa = 0
best_major = ""

for major in uniques:
    
    major_gpas = np.mean(gpa_22_below[majors_22_below == major]) #gets avg gpa for majors_22_below
    
    if major_gpas > max_avg_gpa:
        max_avg_gpa = major_gpas
        best_major = major

print(avg_gpa)
print(max_cs_gpa)
print(top_gpa)
print(older)
print(best_major)


# raise NotImplementedError()
```