### CS 210: Data Management for Data Science
#### Fall 2023

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Arnav Dashaputra"
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## 1. `primes` (10 points)

Write a function to return all the prime numbers less than $n$, where $n$ is a parameter (assume $n>2$). 

``` 
$ primes(40)
[2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37]
``` 

<!-- grade_id: problem-1, points: 0, grade: false, solution: true -->
```python
def primes(n):
    """Returns all prime numbers from 2 to n (excluding n) """
    result = []
    for num in range(2, n):
        
        prime = True
        
        for i in range(2, int(num ** 0.5) + 1):
            if num % i == 0:
                prime = False
                break
                
        if prime:
            result.append(num)
            
    return result
#    raise NotImplementedError()
```

Your function should return `[2, 3, 5, 7]` for $n=10$. Check that it does:

```python
primes(10)
```

Output:
```
[2, 3, 5, 7]
```

<!-- grade_id: problem-1-tests, points: 10, grade: true, solution: false -->
```python
"""Check that primes returns the correct output for several inputs"""
assert primes(3) == [2]
assert primes(6) == [2, 3, 5]
```

---

## 2. `factor` (15 points)

Write a function to return all the prime factors of $n$, where $n$ is a parameter (*Hint: you can call your* `prime` *function from the previous problem*).  

``` 
$ factor(40)
[2, 2, 2, 5]

$ factor(23)
[23]
``` 

<!-- grade_id: problem-2, points: 0, grade: false, solution: true -->
```python
def factor(n):
    """Return all the prime factors of n."""
    factors = []
    for prime in primes(n+1):
        while n % prime == 0:
            factors.append(prime)
            n //= prime #floordiv
    return factors

#    raise NotImplementedError()
```

The factors of `35` are `[5, 7]`. Verify that this is the answer you get:

```python
factor(35)
```

Output:
```
[5, 7]
```

<!-- grade_id: problem-2-tests, points: 15, grade: true, solution: false -->
```python
"""Check that factor returns the correct answer for various inputs."""
assert factor(2) == [2]
assert factor(12) == [2, 2, 3]

```