### CS 210: Data Management for Data Science
#### Fall 2023

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Arnav Dashaputra"
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## 5. `zero_triples` (30 points)

Write a function called `zero_triples` to return all triples in a list of integers that sum to zero. You can assume the list won't contain any duplicates, and a triple should not use the same number more than once. 

```
[In]: zero_triples([1, 2, 4])
[Out]: []

[In]: zero_triples([-3, 1, 4, 2])
[Out]: [[1, 2, -3]]

[In]: zero_triples([-9, 1, -3, 2, 4, 5, -4, -1])
[Out]: [[-9, 4, 5], [1, -3, 2], [-3, 4, -1], [5, -4, -1]]
```


<!-- grade_id: problem-5, points: 0, grade: false, solution: true -->
```python
def zero_triples(ints):
    """Returns all triples in a list of integers that sum to zero. 
    The list of integers is given as a parameter. 
    Assumes the list will not contain any duplicates, and a triple 
    should not use the same number more than once."""

    l = len(ints)
    result = []
    
    for i in range(l):
        for j in range(i + 1, l):
            for k in range(j + 1, l):
                if ints[i] + ints[j] + ints[k] == 0:
                    result.append([ints[i], ints[j], ints[k]]) # adds to end of list
                    
    return result
    
    # raise NotImplementedError()
```

Your function should return `[-3, 1, 2]` for `[-3, 1, 4, 2]`. Check that it does: 

```python
zero_triples([-3, 1, 4, 2])
```

Output:
```
[[-3, 1, 2]]
```

<!-- grade_id: problem-5-tests, points: 30, grade: true, solution: false -->
```python
""" Check that the stats function returns the correct output.  """
assert zero_triples([1, 2, 4])==[]
```