### CS 210: Data Management for Data Science
#### Fall 2023

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Arnav Dashaputra"
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## 3. `increasing_digits` (20 points)

Write a function called `increasing_digits` to count all integers from 1 to $n$ (inclusive) that have all 
digits in increasing order, where $n$ is given as a parameter. For example, 

* $5$ is in increasing order (single digit)
* $237$ is in increasing order ($2<3<7$)
* $227$ is not ($2 \nless 2$)
* $427$ is not ($4 \nless 2$)

```
[In]: increasing_digits(30)
[Out]: 24
```


<!-- grade_id: problem-3, points: 0, grade: false, solution: true -->
```python
def increasing_digits(n):

    count = 0
    
    for i in range(1, n+1):
        val = str(i)
        increasing = True
        
        for j in range(len(val) - 1):
            if val[j] >= val[j+1]:
                increasing = False
                break

        if increasing:
            count += 1
    
    return count
    
    # raise NotImplementedError()
```

Your function should return `9` for $n=10$. Check that it does: 

```python
increasing_digits(10)
```

Output:
```
9
```

<!-- grade_id: problem-3-tests, points: 20, grade: true, solution: false -->
```python
""" Check that the increasing_digits function returns the correct output.  """
assert increasing_digits(5) == 5
```