### CS 210: Data Management for Data Science
#### Fall 2023

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Arnav Dashaputra"
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## 4. `stats` (25 points)

Write a function called `stats` to return the mean, median, and
standard deviation (in that order) of a list of integers, where the list of integers is given as a parameter. 

Suppose the numbers are $x_1,x_2,...,x_n$. Then the mean is $\mu=\frac{1}{n}\Sigma_{i=1}^{n}x_i$. 

The median is the "middle" value is the number of integers ($n$) is odd, or the average of the two "middle" values if $n$ is even (assuming the $x_i$ values are arranged in increasing order). 

The standard deviation is computed as follows: 

$\sigma=\sqrt{\frac{1}{n-1}\Sigma_{i=1}^{n}(x_i-\mu)^2}$.

The stats must be returned in a list. 

Example:

```
[In]: stats([15, 4, 10, 2])
[Out]: [7.75, 7.0, 5.909032633745278]
```


```python
import math
```

<!-- grade_id: problem-4, points: 0, grade: false, solution: true -->
```python
def stats(ints):
    """Returns the mean, median, and standard deviation (in that order) of a 
    list of integers given as a parameter."""
    
    # avg
    mean = sum(ints) / len(ints)
    
    # median
    ordered = sorted(ints)
    mid = len(ints) // 2
    median = ordered[mid]
    if len(ints) % 2 == 0:
        median = (ordered[mid - 1] + ordered[mid]) / 2
    
    # stdev
    std_dev = math.sqrt(sum((x - mean) ** 2 for x in ints) / (len(ints) - 1))    
    
    return [mean, median, std_dev]
```

Your function should return `[7.75, 7.0, 5.909032633745278]` for `[15, 4, 10, 2]`. Check that it does: 

```python
stats([15, 4, 10, 2])
```

Output:
```
[7.75, 7.0, 5.909032633745278]
```

<!-- grade_id: problem-4-tests, points: 25, grade: true, solution: false -->
```python
""" Check that the stats function returns the correct output.  """
assert stats([15, 4, 10, 2]) == [7.75, 7.0, 5.909032633745278]
```